<a href="https://colab.research.google.com/github/Shurui-Zhang/Deep_learning/blob/main/Lab4_2_MLP.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Part 2: Implementing an MLP classifier

```python
# Execute this code block to install dependencies when running on colab
try:
    import torch
except:
    from os.path import exists
    from wheel.pep425tags import get_abbr_impl, get_impl_ver, get_abi_tag
    platform = '{}{}-{}'.format(get_abbr_impl(), get_impl_ver(), get_abi_tag())
    cuda_output = !ldconfig -p|grep cudart.so|sed -e 's/.*\.\([0-9]*\)\.\([0-9]*\)$/cu\1\2/'
    accelerator = cuda_output[0] if exists('/dev/nvidia0') else 'cpu'

    !pip install -q http://download.pytorch.org/whl/{accelerator}/torch-1.0.0-{platform}-linux_x86_64.whl torchvision

try: 
    import torchbearer
except:
    !pip install torchbearer
```

## Getting Started with a Baseline Multi-Layer Perceptron Model

In addition to being a defacto tensor processing and automatic differentiation library, PyTorch provides a general purpose neural network toolbox. Before we start to look at deep convolutional architectures, we can start with something much simpler - a basic multilayer perceptron. Because the MNIST images are relatively small, a fully connected MLP network will have relatively few weights to train; with bigger images, an MLP might not be practical due to the number of weights.

In this section we will create a simple multi-layer perceptron model with a single hidden layer that achieves an error rate of 1.88%. We will use this as a baseline for comparing more complex convolutional neural network models later.

Let's start off by importing the classes and functions we will need.

```python
import torch
import torch.nn.functional as F #包含 torch.nn 库中所有函数 同时包含大量 loss 和 activation function
import torchvision.transforms as transforms
from torch import nn
from torch import optim
from torch.utils.data import DataLoader
from torchvision.datasets import MNIST
```

When developing, it is always a good idea to initialize the random number generator to a constant to ensure that the results of your script are reproducible each time you run it. Once the model is implemented and tested we might remove this to enable variance to be captured over multiple training runs.

```python
# fix random seed for reproducibility
seed = 7
torch.manual_seed(seed)#为CPU设置种子用于生成随机数，以使得结果是确定的
torch.backends.cudnn.deterministic = True #每次返回的卷积算法将是确定的，即默认算法
torch.backends.cudnn.benchmark = False #寻找默认算法
import numpy as np
np.random.seed(seed)
```

## Loading the MNIST data

For the MLP network that we'll shortly define, we need to define a transformation to the data provided by the `MNIST` object that flattens each images into a vector. Before that happens, our transformation must also convert each image from a PIL image to a PyTorch `Tensor`. We can use the classes provided by the `torchvision.transforms` package to compose a set of transforms that given a PIL image will convert the input image to a `Tensor` and reshape that tensor into a vector:

```python
# flatten 28*28 images to a 784 vector for each image
transform = transforms.Compose([
    transforms.ToTensor(),  # convert to tensor
    transforms.Lambda(lambda x: x.view(-1))  # flatten into vector
])
```

Note that for the latter part of the transform we use `Lambda` which allows us to construct a custom transform (rather than use one provided by torchvision). In this case our lambda function returns a new view of the input tensor which is a vector of length equal to the product of the input tensor's dimensions. In this particular case the input tensor had shape `[28, 28]`, so the output will have shape `[784]`. Passing `-1` as the input to the `view` method makes it automatically compute the output shape on the basis of the input.

Now we can load the training and test splits of the MNIST dataset using the torchvision `MNIST` utility class. We can also provide our `transform` object so that it will be applied automatically when the data loads.

```python
trainset = MNIST(".", train=True, download=True, transform=transform)
testset = MNIST(".", train=False, download=True, transform=transform)
```

When we train and evaluate our model, we will require batches of data to be provided. As previously mentioned, the `DataLoader` class can automatically provide batches fetched from a `Dataset` object. The `DataLoader` has the ability to fetch batches in the background (using multiple threads if required) to optimise throughput to the deep network model. It also supports shuffling of the data order (very useful for gradient-based learning!).

```python
# create data loaders
# trainloader = DataLoader(trainset, batch_size=128, shuffle=True, drop_last=True)
# testloader = DataLoader(testset, batch_size=128, shuffle=True, drop_last=True)
trainloader = DataLoader(trainset, batch_size=128, shuffle=True)
testloader = DataLoader(testset, batch_size=128, shuffle=True)
```

DataLoaders are iterable and can be used within a loop to fetch batches of data. Each batch is a tuple containing the images in the first element and the labels in the second. __Complete the following code block to calculate the number of instances of each class by iterating over the training data__:

<!-- grade_id: cell-cd6a3d883391addc, points: 0, grade: false, solution: true -->
```python
class_counts = torch.zeros(10, dtype=torch.int32)
class_labels = torch.tensor([0,1,2,3,4,5,6,7,8,9])
 
for (images, labels) in trainloader:
    for label in labels:
      for i in class_labels:
        if label == class_labels[i]:
          class_counts[i] = class_counts[i] + 1
print(class_counts)
```

Output:
```
tensor([5923, 6742, 5958, 6131, 5842, 5421, 5918, 6265, 5851, 5949],
       dtype=torch.int32)
```

<!-- grade_id: cell-8e6586ecf9f2b748, points: 3, grade: true, solution: false -->
```python
assert class_counts.sum()==60000
```

__Answer the following questions (enter the answer in the box below each one):__

__1.__ Do all of the batches have the same size?

<!-- grade_id: cell-d7476483eabfb245, points: 1, grade: true, solution: true -->
出最后一个batche外，其它batches的size都一样，都为128，最后一个为96。但每个class中包含的实例个数不同

__2.__ How can we configure `DataLoader` to control this?

<!-- grade_id: cell-c87b3cc38b91cb03, points: 1, grade: true, solution: true -->
可以通过设置dataloader的drop_last属性为True将最后的batch弃用。

## Defining the MLP Model

We are now ready to create our simple neural network model. We will define our model in a class that extends `nn.Module`. `nn.Module` subclasses must do a minimum of one thing: implement the `forward` method which takes a batch of data and performs the forward-pass. PyTorch's autograd system will take care of computing the gradients of the forward pass for us. In the code below we'll also make use of the constructor of our model to instantiate the hidden and output layers.

```python
# define baseline model
class BaselineModel(nn.Module):
    def __init__(self, input_size, hidden_size, num_classes):
        super(BaselineModel, self).__init__() #调用父类构造函数
        self.fc1 = nn.Linear(input_size, hidden_size) #nn.Linear()是用于设置网络中的全连接层的。定义fc1（fullconnect）全连接函数1为线性函数：y = Wx + b，把input_size个节点连接到hidden_size个节点上
        self.fc2 = nn.Linear(hidden_size, num_classes)  
    
    def forward(self, x):
        out = self.fc1(x) #让输入x经过全连接1
        out = F.relu(out) #用relu做为第一层激活函数
        out = self.fc2(out)
        if not self.training:
            out = F.softmax(out, dim=1)
        return out
```

The model is a simple neural network with one hidden layer. A rectifier linear unit (ReLU) activation function is used for the neurons in the hidden layer. Note how we use the PyTorch 'Functional API' for stateless operations like applying the ReLU; we could construct a `torch.nn.ReLU` in the constructor and save it in an instance variable instead, but ultimately this would lead to more lines of code.

The `nn.Module` class defines a instance variable called `training` that is set to `True` when the model is being trained and `False` when it is being evaluated after being trained. In our model definition we've used a softmax activation function on the output layer to turn the outputs into probability-like values, but have only set this to be enabled when we are not training the model. We've done this because we will use PyTorch's implementation of Cross Entropy Loss (`nn.CrossEntropyLoss`) during training which implicitly adds a softmax before a logarithmic loss (technically it adds a log-softmax activation followed by a negative log-likelihood loss as this has much better numerical properties)（只在测试时才对输出层进行softmax的原因）. If you've used a library like Keras before you should note that this approach is different to the one taken there - Keras expects you to explicitly add a `Softmax` layer to use with its `categorical_crossentropy` loss.

In our case the softmax isn't actually necessary for model evaluation if we're only interested in the most likely class; the _logits_ (unscaled log probabilities) provided by the final fully connected layer before the softmax can be used directly as the largest logit will correspond to the most likely class.

## Training and Evaluating the Model
We can now fit and evaluate the model. One of the design decisions of PyTorch is that everything should be explicit so we have full control over our models and the training process. This means that we actually need to write the model training loop by hand, and perform each of the various operations (perform the forward-pass, compute the loss, perform the backward-pass, and update the weights). In the code below we'll fit the model to the data over 10 epochs using batches of 128 images provided by the `DataLoader` defined previously. We'll make use of the ADAM optimiser as it broadly tends to work well practically despite its limitations. Running the code below should take just a couple of minutes to complete the training.

```python
# build the model 
model = BaselineModel(784, 784, 10)

# define the loss function and the optimiser
loss_function = nn.CrossEntropyLoss()
optimiser = optim.Adam(model.parameters()) #model.parameters()获取module中的参数

# the epoch loop
for epoch in range(10):
    running_loss = 0.0
    for data in trainloader:
        # get the inputs. data是一个元组，包含图片和label
        inputs, labels = data

        # zero the parameter gradients
        optimiser.zero_grad()

        # forward + loss + backward + optimise (update weights)
        outputs = model(inputs) #调用forward函数：model(x)调用基类的__call__方法，__call__再调用forward
        loss = loss_function(outputs, labels)
        loss.backward()
        optimiser.step()

        # keep track of the loss this epoch
        running_loss += loss.item()
    print("Epoch %d, loss %4.2f" % (epoch, running_loss))
print('**** Finished Training ****')
```

Output:
```
Epoch 0, loss 134.37
Epoch 1, loss 51.74
Epoch 2, loss 33.02
Epoch 3, loss 23.58
Epoch 4, loss 16.89
Epoch 5, loss 12.29
Epoch 6, loss 9.90
Epoch 7, loss 7.14
Epoch 8, loss 5.36
Epoch 9, loss 3.42
**** Finished Training ****
```

In the above we added a crude indicator of progress by printing the total loss at the end of each epoch. Hopefully the loss went down as the model learned! 

At this point it would be good to compute the overall accuracy of the test set. __Use the following code block to finish implementation of the accuracy computation.__ Note that before the code you'll implement we've made a call to `model.eval()` - this sets the model into evaluation mode and supresses non-training things (gradients, and things such as dropout being applied/computed).

<!-- grade_id: cell-13268e61111fdb7e, points: 5, grade: true, solution: true -->
```python
model.eval()

# Compute the model accuracy on the test set
correct = 0
total = 0

for data in testloader:
  inputs, labels = data
  outputs = model(inputs)
  loss = loss_function(outputs, labels)
  for i in range(outputs.shape[0]):
    list = outputs[i].tolist()
    if list.index(max(list)) == labels[i].item():
      correct = correct + 1
    total = total + 1

print('Test Accuracy: %2.2f %%' % ((100.0 * correct) / total))

```

Output:
```
tensor(1.4834, grad_fn=<NllLossBackward>)
tensor(1.4917, grad_fn=<NllLossBackward>)
tensor(1.4621, grad_fn=<NllLossBackward>)
tensor(1.4638, grad_fn=<NllLossBackward>)
tensor(1.4912, grad_fn=<NllLossBackward>)
tensor(1.4703, grad_fn=<NllLossBackward>)
tensor(1.4827, grad_fn=<NllLossBackward>)
tensor(1.4738, grad_fn=<NllLossBackward>)
tensor(1.4769, grad_fn=<NllLossBackward>)
tensor(1.4912, grad_fn=<NllLossBackward>)
tensor(1.4659, grad_fn=<NllLossBackward>)
tensor(1.4762, grad_fn=<NllLossBackward>)
tensor(1.4700, grad_fn=<NllLossBackward>)
tensor(1.5027, grad_fn=<NllLossBackward>)
tensor(1.4871, grad_fn=<NllLossBackward>)
tensor(1.4995, grad_fn=<NllLossBackward>)
tensor(1.5020, grad_fn=<NllLossBackward>)
tensor(1.4797, grad_fn=<NllLossBackward>)
tensor(1.4789, grad_fn=<NllLossBackward>)
tensor(1.4808, grad_fn=<NllLossBackward>)
tensor(1.4661, grad_fn=<NllLossBackward>)
tensor(1.4887, grad_fn=<NllLossBackward>)
tensor(1.4766, grad_fn=<NllLossBackward>)
tensor(1.4945, grad_fn=<NllLossBac
```

__In multi-class classification tasks it is often instructive to explore the accuracy of each class. Use the code block below to complete the function that produces the per-class accuracy:__

<!-- grade_id: cell-64fb2d2effc2543d, points: 5, grade: true, solution: true -->
```python
# Compute the model accuracy on the test set
class_correct = torch.zeros(10)

class_total = torch.zeros(10)
class_labels = torch.tensor([0,1,2,3,4,5,6,7,8,9])
 
for (images, labels) in testloader:
    for label in labels:
      for i in class_labels:
        if label == class_labels[i]:
          class_total[i] = class_total[i] + 1

for data in testloader:
  inputs, labels = data
  outputs = model(inputs)
  for i in range(outputs.shape[0]):
    list = outputs[i].tolist()
    if list.index(max(list)) == labels[i].item():
      class_correct[labels[i].item()] = class_correct[labels[i].item()] + 1

for i in range(10):
    print('Class %d accuracy: %2.2f %%' % (i, 100.0*class_correct[i] / class_total[i]))
```

Output:
```
Class 0 accuracy: 99.08 %
Class 1 accuracy: 99.56 %
Class 2 accuracy: 98.26 %
Class 3 accuracy: 97.82 %
Class 4 accuracy: 97.45 %
Class 5 accuracy: 97.53 %
Class 6 accuracy: 98.23 %
Class 7 accuracy: 98.44 %
Class 8 accuracy: 97.43 %
Class 9 accuracy: 97.92 %
```

```python

```